## Instructions
- See deadline on the course web page
- This problem set is solved individually. See examination rules on the course web page and the explanation of the examination procedure below.
- The two notebooks for each problem set contain a number of basic and extra problems; you can choose which and how many to work on. The extra problems are usually more challenging.
- Students are allowed to discuss together and help each other when solving the problems. However, every student must understand their submitted solution in the sense that they should be able to explain and discuss them with a peer or with a teacher.
- While discussions with your peers are allowed (and even encouraged), direct plagiarism is not. Every student must reach their own understanding of submitted solutions according to the definition in the previous point.
- The use of coding assistance from code generating artificial intelligence tools is allowed. However, every student must reach their own understanding of submitted solutions (including employed algorithms) according to the definition above.
- Some problems include checkpoints in the form of `assert` statements. These usually check some basic functionality and you should make sure that your code passes these statements without raising an `AssertionError`. 
- Do not use other python modules than the ones included in the `environment.yml` file in the course github repo. 

- **Important:** The grading of problem sets requires **all** of the following actions:
  1. Make sure to always complete **Task 0** in the header part of the notebook and that this part does not raise any `AssertionError`(s).
  1. **Complete** the corresponding questions in Yata for every task that you have completed. This usually involves copying and pasting some code from your solution notebook and passing the code tests. You need to have a green check mark on Yata to get the corresponding points.
  1. **Upload** your solution in the form of your edited version of this Jupyter notebook via the appropriate assignment module in Canvas (separate for basic and extra tasks). It is the code and results in your submitted notebook that is considered to be your hand-in solution.
  1. If selected, be **available for a discussion** of your solution with one of the teachers on the Monday afternoon exercise session directly following the problem set deadline. No extra preparation is needed for these discussions apart from familiarity with your own solution. A list of randomly selected students will be published on the course web page around Monday noon. During the afternoon session that same day, students will be called in the numbered order until the end of the list (or the end of the exercise session). You must inform the responsible teacher as soon as possible following the publication of the student list if you can not be physically present at the exercise session (in which case we will have the discussion on zoom). An oral examination (on all aspects of the course) will be arranged during the exam week for students that do not show up for their discussion slot, or that fail to demonstrate familiarity with their hand-in solutions.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

- Make sure that the **run time is smaller than a few minutes**. If needed you might have to reduce some computational tasks; e.g. by decreasing the number of grid points or sampling steps. Please ask the supervisors if you are uncertain about the run time. 

- Your solutions are usually expected where it says `YOUR CODE HERE` or <font color="red">"PLEASE WRITE YOUR ANSWER HERE"</font>.

### Task 0 
#### (0 points)

By changing the below boolean variable `student_self_assessment` to `True` you attest that:
- All handed in solutions were produced by yourself in the sense that you understand your solutions and should be able to explain and discuss them with a peer or with a teacher.


<!-- grade_id: student-self-assessment, points: 0, grade: false, solution: true -->
```python
student_self_assessment = False

# 
# YOUR CODE HERE
# 
```

<!-- grade_id: cell-795bedd2908899aa, points: 0, grade: true, solution: false -->
```python
assert student_self_assessment == True, 'You must assert the individual solution statements.'
```

# Problem Set 1
## Basic problems
### Learning from data [TIF285], Chalmers, Fall 2023

Last revised: 18-Aug-2023 by Christian Forssen [christian.forssen@chalmers.se]

## Problem 1 (2 points)

### Installations
Perform the installations and preparations that are described in the Getting Started instructions. At the end you should have:

1. downloaded the current version of the course material from the github repository or from the course web page;
2. a running python installation that includes the modules listed in the environment.yml file (e.g. numpy, matplotlib, pandas, emcee, scikit-learn, ...);
3. been able to open and run the Jupyter Notebooks with the first week exercises.
Ask the computer lab supervisors for assistance if needed.

<!-- grade_id: import, points: 0, grade: false, solution: true -->
```python
# Modules needed for tests and file system operations
import sys

import os

# Where to save the figures and data files
DATA_ID = "DataFiles/"
if not os.path.exists(DATA_ID):
    os.makedirs(DATA_ID)

# Make sure that you are running python with version >= 3.x
#
# Import the following python modules with
# the specified abreviations:
# ---
# numpy as np
# scipy as scipy
# scipy.stats as stats
# pandas as pd
# matplotlib.pyplot as plt
# sklearn as skl
# emcee as emcee


# 
# YOUR CODE HERE
# 
```

<!-- grade_id: correct_import, points: 1, grade: true, solution: false -->
```python
# TESTS TO CHECK YOUR CODE
# ---
assert sys.version_info.major>=3, \
    'You are running Python version'+\
    f'{sys.version_info.major}.{sys.version_info.minor}'

modules = [('numpy','np'), ('scipy', 'scipy'), \
           ('pandas', 'pd'), ('matplotlib.pyplot', 'plt'), \
           ('sklearn', 'skl'), ('emcee', 'emcee')]
for (_module, _module_abbrev) in modules:
    assert _module in sys.modules and _module_abbrev in dir(),\
        f'Module {_module} not loaded properly.'
```

### Random samples
In this course there are $N$ students and three problem sets.
The random selection of students for problem set discussions is performed by drawing $N$ samples of three random variables $(X_1, X_2, X_3)$. These are arranged in an array of shape (N,3) as follows 
$$
\left(
\begin{array}{ccc}
x_1(1) & x_2(1) & x_3(1) \\
x_1(2) & x_2(2) & x_3(2) \\
\vdots & \vdots & \vdots \\
x_1(N) & x_2(N) & x_3(N)
\end{array}
\right).
$$
The students with the lowest values of the respective random variable will then be invited for discussions with a teaching assistant. E.g., for problem set 1 there will be an ordered list of students $(k, l, m, \ldots)$, where $x_1(k) < x_1(l) < x_1(m) < \ldots$.

It is considered fair that all students have the same probability to draw at least one small number, but it should be less likely to draw two small numbers. This can be achieved by making the random variables dependent, more specifically by making the covariance negative (anticorrelated). Here we use a correlated, multivariate Gaussian
$$
(X_1, X_2, X_3) \sim \mathcal{N}\left( \boldsymbol{\mu}, \boldsymbol{\Sigma} \right),
$$
with mean and covariance
$$
\boldsymbol{\mu} = (0,0,0),\quad
\boldsymbol{\Sigma} = \left(
\begin{array}{ccc}
1 & -0.49 & -0.49 \\
-0.49 & 1 & -0.49 \\
-0.49 & -0.49 & 1
\end{array}
\right),
$$ 
which means that every univariate, marginal distribution will be a standard normal ($\mu=0, \sigma^2=1$).

Take a minute to consider the reason for making the random variables anticorrelated (it is instructive to visualize the multivariate distribution) and the fact that the correlation factor cannot be more negative than -0.5.

**Task (see also Yata):** Consider a class with 6 students. Use the method `multivariate_normal` from `scipy.stats` and generate an array of shape (6,3) with samples from the distribution defined in Problem 1. Print the result. 

*Important*: Use the random seed 2023 when generating the samples.

<!-- grade_id: P1array, points: 0, grade: false, solution: true -->
```python
# 
# YOUR CODE HERE
# 
```

**Task (see also Yata):** Create a `Pandas.DataFrame` with the sample array from the previous task. The columns should be labeled ['X1', 'X2', 'X3'] while the row indices should be ['Student 1', 'Student 2', 'Student 3', ...]. Print the DataFrame (in the notebook you can try the more fancy option `display`, while you should use `print` on Yata).

Assume that only the first two students from each ordered list are selected. Which student(s) will not be selected for any of the problem sets.

<!-- grade_id: P1df, points: 0, grade: false, solution: true -->
```python
# 
# YOUR CODE HERE
# 
```

## Problem 2 (3 points)

**Generate data**

<!-- grade_id: P2generate, points: 0, grade: false, solution: true -->
```python
# Generate noisy data with a quadratic feature
# use the following code:
np.random.seed(42)
m = 100 # Number of data

# X are picked uniform random [0,2]
X = 2 * np.random.rand(m, 1)
# Linear relation to the predicted value, but with Gaussian noise (mean=0, standard deviation=0.2)
theta_true = [0.25, 1, 0.75]
noise_std = 0.2
y = theta_true[2] * X**2 + theta_true[1] * X + theta_true[0] +  np.random.normal(loc=0.0, scale=noise_std, size=(m,1))

# Below is a hidden code block that is used in the solution notebook to save the data. 
# 
# Please ignore the comment in this cell that says "YOUR CODE HERE". It gets added automatically.
# No solution code is needed here.
#
# 
# YOUR CODE HERE
# 
```

### (a) Linear regression using the Normal Equation

**Task (see also Yata):** Create a python function `design_matrix` that returns the design matrix for a polynomial model. 

<!-- grade_id: design_matrix, points: 0, grade: false, solution: true -->
```python
def design_matrix(X, degree=2):
    """
    Returns a design matrix.
    
    Args:
        X: Array of shape (m,1) with 'm' independent data.
        degree: Integer with the degree of the polynomial. 
                  Note that a degree-n polynomial has n+1 coefficients.
                  
    Returns:
        X_d: Design matrix of shape (m, order+1).
    """
    # 
    # YOUR CODE HERE
    # 
```

<!-- grade_id: correct_design_matrix, points: 1, grade: true, solution: false -->
```python
# TESTS TO CHECK YOUR CODE
# ---
assert design_matrix(X).shape == (len(X),3)
assert design_matrix(X)[:,0].all() == 1
assert design_matrix(X)[0,1] == X[0]
assert design_matrix(X)[0,2] == X[0]**2

ratio = design_matrix(X)[:,2]/design_matrix(X)[:,1]**2
assert ratio.all() == 1
```

**Task (see also Yata):** Complete the python method `solve_normal_equation` that solves the normal equation using matrix inversion. This function can be used to fit a quadratic model to the data generated for this problem.  

<!-- grade_id: solve_normal_equation, points: 0, grade: false, solution: true -->
```python
def solve_normal_equation(X_d, y):
    """
    Solve the normal equation.
    
    Args:
        X_d: Design matrix of shape (m,n) with 'm' independent data
               and 'n' features.
        y: Dependent data of shape (m,1).
                  
    Returns:
        theta_best: Best parameters, array of shape (n,).
    """
    # 
    # YOUR CODE HERE
    # 
```

<!-- grade_id: correct_solve_normal_equation, points: 1, grade: true, solution: false -->
```python
# TESTS TO CHECK YOUR CODE
# ---
assert (solve_normal_equation(design_matrix(X), y)).shape==(3,),\
    'Return object has wrong shape. Maybe the `flatten` method will be useful?'
theta_star = np.array([0.36228054, 0.72777588, 0.86621274])
theta_residuals = solve_normal_equation(design_matrix(X), y) - theta_star
assert (abs(theta_residuals)<1e-7).all(), f'The solution to the normal equation is of low precision. theta_residulas = {abs(theta_residuals)}'
```

### (b) Gradient descent
Gradient descent optimization is not really needed for linear regression since we can solve the normal equation exactly. However, it is instructive to implement it yourself and to be able to compare with a known solution.

**Task (see also Yata)** 
Define a function to perform gradient descent optimization on a linear model that is defined by a design matrix and the data (see further instructions in the code cell).

<!-- grade_id: gradient_descent_function, points: 0, grade: false, solution: true -->
```python
def gradient_descent_linear_regression(X_d, y, theta_start, eta=0.1, n_iterations=1000):
    """
    Find optimized parameters in linear regression using (batch) gradient descent.
    
    Args:
        X_d: Design matrix of shape (m,n) with 'm' independent data
               and 'n' features (ndarray).
        y: Dependent data of size m (ndarray).
        theta_start: List-like initial guess for parameters (size n) 
        eta: learning rate (default 0.1) (float)
        n_iterations: Number of iterations (epochs) (default 1000) (integer)
                  
    Returns:
        theta_optimum: Optimized parameters, array of shape (n,).
    """
    (m,n) = X_d.shape
    assert len(y)==m, f'The size of y must be {m}. It is: {len(y)}.'
    y = y.reshape(-1,1) # Turn into a (m,1) array.
    assert len(theta_start)==n, f'The initial guess must be of size {n}. It is: {len(theta_start)}.'
    
    # 
    # YOUR CODE HERE
    # 
```

<!-- grade_id: correct_gradient_descent_function, points: 1, grade: true, solution: false -->
```python
# TESTS TO CHECK YOUR CODE
# ---
assert (gradient_descent_linear_regression(design_matrix(X), y,[0,0,0],n_iterations=1)).shape==(3,),\
    'Return object has wrong shape. Maybe the `flatten` method will be useful?'
X_d = np.c_[np.ones((m, 1)), X, X**2]
# Solution to normal equation
theta_star = np.array([0.36228054, 0.72777588, 0.86621274])
# Check gradient
gradient_optimum = theta_star - gradient_descent_linear_regression(X_d, y, theta_star,eta=1.0,n_iterations=1)
assert (abs(gradient_optimum)<1e-7).all(), 'Gradient at optimum is not zero (elements > 10**-7)'
# Check final optimum (tolerance 10**-2)
theta_residuals = gradient_descent_linear_regression(X_d, y, [0,0,0]) - theta_star
assert (abs(theta_residuals)<1e-2).all(), 'Residuals of GD optimum are too large (elements > 10**-2)'
```

### (c) Comparison of solvers

**Tasks (see also Yata)** 
Apply both your solver of the normal equation and your gradient descent optimizer to the same data and quadratic model that was introduced at the beginning of this problem.
- Print and compare the coefficients from (i) the true polynomial data generator to the linear regression fits to the (noisy) data from both (ii) the exact solution of the normal equation and (iii) the result of the gradient descent optimization. 
- Plot the data and the model predictions in the same figure.

<!-- grade_id: cell-fbd1f3bdda876913, points: 0, grade: false, solution: true -->
```python
# 1. Print and compare the coefficients from 
#   (i) the true data generator, 
#   (ii) the solution of the normal equation, 
#   (iii) the optimum from gradient descent.
#
# 2. Plot the data, the true model and the fit model predictions in the same figure.

# 
# YOUR CODE HERE
# 
```

## Problem 3 (3 points)

There are three files in the directory `DataFiles`:
- `dataset1.dat`
- `dataset2.dat`
- `dataset3.dat`

Each data files contains two columns. The first column corresponds to the independent variables (the array X), and the second column corresponds to the dependent ones (the array y).

<!-- grade_id: cell-a51f823aed196b24, points: 0, grade: false, solution: true -->
```python
# This cell is used in the solution notebook to generate and save the data. 
# It is hidden in the student version.
# 
# Please ignore the comment in this cell that says "YOUR CODE HERE". It gets added automatically.
# No solution code is needed here.
# ---
# 
# YOUR CODE HERE
# 
```

### (a) Implement linear regression and the MSE cost function for training and validation
**Tasks (see also Yata for the final numerical comparison)** 
- Load a data set and split it into 60% training and 40% validation data using the python function below.
- Implement a linear regression function that takes training data as input and returns a best-fit parameter vector for a polynomial model of a specified degree (see code template below).
- Implement a cost function that takes data and model parameters as input and returns the mean-squared error (see code template below).
- Apply this functionality to one of the datasets (see Yata for numerical checks).

Hint: use the `design_matrix` and `solve_normal_equation` methods from Problem 2. 

```python
# built-in convenience function for splitting data
from sklearn.model_selection import train_test_split

def load_data(datafile, train_size=0.6):
    """
    Reads data from file and returns training and validation sets.
    
    Args:
        datafile: String with data filename path. The data file 
            should contain two columns: x, y
        train_size: float indicating the fraction of training data
            (default: 0.6)
            
    Returns:
        (X_train, X_val, y_train, y_val): Tuple with four arrays 
            with training and validation data.
    """
    X, y = np.loadtxt(datafile, unpack=True)
    m = len(X)
    X = X.reshape(m,1); y = y.reshape(m,1)

    X_train, X_val, y_train, y_val = \
        train_test_split(X, y, train_size=train_size, random_state=42)
    return (X_train, X_val, y_train, y_val)
```

<!-- grade_id: cell-aa16bd9e1a61d370, points: 0, grade: false, solution: true -->
```python
# Implement a linear regression function that takes 
# training data as input and returns a best-fit parameter 
# vector for a polynomial model of a specified degree.
def linear_regression(X, y, degree=2):
    """
    Performs linear regression for a polynomial model.
    
    Args:
        X: Array of shape (m,1) with 'm' independent data.
        y: Array of shape (m,1) with 'm' dependent data.
        degree: Integer with the degree of the polynomial. 
                  Note that a degree-n polynomial has n+1 coefficients.
                  
    Returns:
        theta_fit: Best fit parameters. Array of shape (degree+1,)
    """
    # 
    # YOUR CODE HERE
    # 
```

<!-- grade_id: cell-b26c3f6a60009e39, points: 1, grade: true, solution: false -->
```python
# TESTS TO CHECK YOUR CODE
# ---
datafile = f'{DATA_ID}/PS1_Prob3_dataset1.txt'
(X_train, X_val, y_train, y_val) = load_data(datafile)
assert len(linear_regression(X_train, y_train, degree=3))==4
a0_degree3fit_residual = \
    linear_regression(X_train, y_train, degree=3)[0] - 0.04772979
assert abs(a0_degree3fit_residual) < 1e-5

datafile = f'{DATA_ID}/PS1_Prob3_dataset3.txt'
(X_train, X_val, y_train, y_val) = load_data(datafile)
theta_star_degree4fit = np.array([ 1.42786627,  2.92656911,  0.9522876 , \
                                  -0.38212683, -0.00978446])
theta_residuals = linear_regression(X_train, y_train, degree=4).reshape(-1,) \
    - theta_star_degree4fit
assert (abs(theta_residuals)<1e-5).all(), print('theta_residuals=\n',theta_residuals)
```

<!-- grade_id: MSE, points: 0, grade: false, solution: true -->
```python
# Implement a cost function that takes data and polynomial model 
# parameters as input and returns the mean-squared error.
def mean_squared_error(X, y, theta):
    """
    Compute the mean-squared error for data and a polynomial fit.
    
    Args:
        X: Array of shape (m,1) with 'm' independent data.
        y: Array of shape (m,1) with 'm' dependent data.
        theta: Parameter array [shape (degree+1,)]. 
            The ordering corresponds to the constant term first.
            
    Return:
        MSE (float): Mean-squared error defined as
            MSE = (1/m) * sum_i (y[i] - y_model[i])**2,
            where y_model[i] = \sum_m theta[m]*X[i]**m 
    """
    # 
    # YOUR CODE HERE
    # 
```

```python
# TESTS TO CHECK YOUR CODE
# ---
X_train=np.array([[1],[2],[3]])
y_train=np.array([[2],[5],[10]])
assert mean_squared_error(X_train, y_train, np.array([1,0,1]))==0
assert mean_squared_error(X_train, y_train, np.array([0,0,0]))==43
```

<!-- grade_id: polynomial_regression, points: 0, grade: false, solution: true -->
```python
# Implement a function that takes training and validation data plus 
# polynomial model degree as input, performs the optimization to 
# the training data and returns the mean-squared error for both training 
# and validation data as well as the best fit parameters.
def polynomial_regression( data, degree):
    """
    Compute the mean-squared error for data and a polynomial fit.
    
    Args:
        data = (X_train, X_val, y_train, y_val): Tuple with four arrays 
            with training and validation data.
        degree: Integer with the degree of the polynomial. 
                  Note that a degree-n polynomial has n+1 coefficients.
            
    Return:
        MSE_train: Mean-squared error of training data
        MSE_val: Mean-squared error of validation data
        theta_fit: Best fit parameters [array of shape (degree+1,)]
    """
    # 
    # YOUR CODE HERE
    # 
```

<!-- grade_id: correct_MSE, points: 1, grade: true, solution: false -->
```python
# TESTS TO CHECK YOUR CODE
# ---
X_val = np.array([[4],[5]])
y_val = np.array([[17],[27]])
MSE_train, MSE_val, theta_fit = \
    polynomial_regression( (X_train, X_val, y_train, y_val), 2)
assert MSE_val-0.5 < 1e-5
```

```python
### Final, numerical results. 
# See also Yata
datafile = f'{DATA_ID}/PS1_Prob3_dataset3.txt'
(X_train, X_val, y_train, y_val) = load_data(datafile)
degrees = range(6)
print('Model degree  MSE(train)    MSE(val)')
for degree in degrees:
    MSE_train, MSE_val, theta_fit = \
        polynomial_regression( (X_train, X_val, y_train, y_val), degree)
    print(f'{degree:>12}{"":5}{MSE_train:7.3f}{"":5}{MSE_val:7.3f}  {theta_fit[0]} {theta_fit[-1]}')
```

```python
# TESTS TO CHECK YOUR CODE
# ---
datafile = f'{DATA_ID}/PS1_Prob3_dataset3.txt'
(X_train, X_val, y_train, y_val) = load_data(datafile)
MSE_train, MSE_val, theta_fit = \
    polynomial_regression( (X_train, X_val, y_train, y_val), 4)
assert abs(MSE_train-3.8740275553091648)<1e-5
assert abs(MSE_val-2.2031653660595256)<1e-5
assert abs(theta_fit[0]-1.4278662700653606)<1e-5
```

### (b) Train and validate linear regression with different polynomial models
**Tasks (see also Yata)** 
- For each data set you should then perform linear regression using polynomial models of order 1,2,3,4,5, and 20.
- Finally, print the fit coefficients for each polynomial model that was considered and print also the mean-squared error (MSE) for both the training and the validation sets.

<!-- grade_id: cell-49fbcaec1138af05, points: 0, grade: false, solution: true -->
```python
# 
# YOUR CODE HERE
# 
```

### (c) Concluding discussion
**Tasks (see also Yata)** 
Answer the following two questions:
- Which degrees of polynomials do you think was used for generating the different datasets?
- Which dataset do you think has the most noise?

You should be able to discuss your reasoning.

<!-- grade_id: cell-93dac302fd63f96a, points: 1, grade: true, solution: true -->
```python
# This cell is used only in the solution notebook.
# 
# Please ignore the comment in this cell that says "YOUR CODE HERE". It gets added automatically.
# No solution code is needed here.
# ---
# 
# YOUR CODE HERE
# 
```

## Problem 4 (2 points)

### (a) Applying Bayesian rules of probability to a standard medical example

Suppose there is an unknown disease (call it UD) that does not give any symptoms in its early phase, but that there is a test for it.

a. The false positive rate is 2.3%. ("False positive" means the test says you have UD, but you don't.) <br>
b. The false negative rate is 1.4%. ("False negative" means you have UD, but the test says you don't.)

Assume that 1 in 10,000 people have the disease. You are given the test and get a positive result.  Your ultimate goal is to find the probability that you actually have the disease. 
$% Some LaTeX definitions we'll use.
\newcommand{\pr}{\textrm{p}}
$

We'll do it using the Bayesian rules.

We'll use the notation:

* $H$ = "you have UD"
* $\overline H$ = "you do not have UD"  
* $D$ = "you test positive for UD"
* $\overline D$ = "you test negative for UD"  

**Tasks (see also Yata)** 

Answer the following questions. Some of them are repeated on Yata.

a. *Before doing a calculation (or thinking too hard :), does your intuition tell you the probability you have the disease is high or low?*
<br>

b. *In the $p(\cdot | \cdot)$ notation, what is your ultimate goal?*
<br>

c. *Express the false positive rate in $p(\cdot | \cdot)$ notation.* \[Ask yourself first: what is to the left of the bar?\]
<br>

d. *Express the false negative rate in $p(\cdot | \cdot)$ notation. By applying the sum rule, what do you also know? (If you get stuck answering the question, do the next part first.)* 
<br>

e. *Should $p(D|H) + p(D|\overline H) = 1$?
    Should $p(D|H) + p(\overline D |H) = 1$?
    (Hint: does the sum rule apply on the left or right of the $|$?)*
<br>

f. *Apply Bayes' theorem to your result for your ultimate goal (don't put in numbers yet).
   What other probabilities do we need?*
<br>

<!-- grade_id: bayes_medical, points: 0, grade: false, solution: true -->
```python
# This cell is used only in the solution notebook.
# 
# Please ignore the comment in this cell that says "YOUR CODE HERE". It gets added automatically.
# No solution code is needed here.
# ---
# 
# YOUR CODE HERE
# 

```

### (b) Applying Bayesian rules of probability to a standard medical example (cont.)

**Tasks (see also Yata)**

<!-- grade_id: medical_example, points: 0, grade: false, solution: true -->
```python
# Please fill the probabilities as values for the 
# corresponding keys in the following dictionary.
medical_example_probabilities = {}
medical_example_probabilities['p(D|Hbar)'] = 0.0
medical_example_probabilities['p(Dbar|H)'] = 0.0
medical_example_probabilities['p(D|H)'] = 0.0
medical_example_probabilities['p(H,Hbar|D)'] = 0.0
medical_example_probabilities['p(Hbar)'] = 0.0
medical_example_probabilities['p(D)'] = 0.0
medical_example_probabilities['p(H|D)'] = 0.0

# 
# YOUR CODE HERE
# 
```

<!-- grade_id: correct_1_medical_example, points: 1, grade: true, solution: false -->
```python
for key in ['p(D|Hbar)', 'p(Dbar|H)', 'p(D|H)']:
    assert medical_example_probabilities[key] > 0.
    assert medical_example_probabilities[key] < 1.
    
assert medical_example_probabilities['p(H,Hbar|D)'] <= 1.0
```

<!-- grade_id: correct_2_medical_example, points: 1, grade: true, solution: false -->
```python
for key in ['p(Hbar)', 'p(D)', 'p(H|D)']:
    assert medical_example_probabilities[key] > 0.
    assert medical_example_probabilities[key] < 1.
```