## Instructions
- See deadline on the course web page
- This problem set is solved individually. See examination rules on the course web page and the explanation of the examination procedure below.
- The two notebooks for each problem set contain a number of basic and extra problems; you can choose which and how many to work on. The extra problems are usually more challenging.
- Students are allowed to discuss together and help each other when solving the problems. However, every student must understand their submitted solution in the sense that they should be able to explain and discuss them with a peer or with a teacher.
- While discussions with your peers are allowed (and even encouraged), direct plagiarism is not. Every student must reach their own understanding of submitted solutions according to the definition in the previous point.
- The use of coding assistance from code generating artificial intelligence tools is allowed. However, every student must reach their own understanding of submitted solutions (including employed algorithms) according to the definition above.
- Some problems include checkpoints in the form of `assert` statements. These usually check some basic functionality and you should make sure that your code passes these statements without raising an `AssertionError`. 
- Do not use other python modules than the ones included in the `environment.yml` file in the course github repo. 

- **Important:** The grading of problem sets requires **all** of the following actions:
  1. Make sure to always complete **Task 0** in the header part of the notebook and that this part does not raise any `AssertionError`(s).
  1. **Complete** the corresponding questions in Yata for every task that you have completed. This usually involves copying and pasting some code from your solution notebook and passing the code tests. You need to have a green check mark on Yata to get the corresponding points.
  1. **Upload** your solution in the form of your edited version of this Jupyter notebook via the appropriate assignment module in Canvas (separate for basic and extra tasks). It is the code and results in your submitted notebook that is considered to be your hand-in solution.
  1. If selected, be **available for a discussion** of your solution with one of the teachers on the Monday afternoon exercise session directly following the problem set deadline. No extra preparation is needed for these discussions apart from familiarity with your own solution. A list of randomly selected students will be published on the course web page around Monday noon. During the afternoon session that same day, students will be called in the numbered order until the end of the list (or the end of the exercise session). You must inform the responsible teacher as soon as possible following the publication of the student list if you can not be physically present at the exercise session (in which case we will have the discussion on zoom). An oral examination (on all aspects of the course) will be arranged during the exam week for students that do not show up for their discussion slot, or that fail to demonstrate familiarity with their hand-in solutions.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

- Make sure that the **run time is smaller than a few minutes**. If needed you might have to reduce some computational tasks; e.g. by decreasing the number of grid points or sampling steps. Please ask the supervisors if you are uncertain about the run time. 

- Your solutions are usually expected where it says `YOUR CODE HERE` or <font color="red">"PLEASE WRITE YOUR ANSWER HERE"</font>.

### Task 0 
#### (0 points)

By changing the below boolean variable `student_self_assessment` to `True` you attest that:
- All handed in solutions were produced by yourself in the sense that you understand your solutions and should be able to explain and discuss them with a peer or with a teacher.


<!-- grade_id: student-self-assessment, points: 0, grade: false, solution: true -->
```python
student_self_assessment = False

# 
# YOUR CODE HERE
# 
```

<!-- grade_id: cell-795bedd2908899aa, points: 0, grade: true, solution: false -->
```python
assert student_self_assessment == True, 'You must assert the individual solution statements.'
```

# Problem Set 1
## Extra problems
### Learning from data [TIF285], Chalmers, Fall 2023

Last revised: 20-Aug-2023 by Christian Forssén [christian.forssen@chalmers.se]

<!-- grade_id: cell-30c16b2db381306f, points: 0, grade: false, solution: true -->
```python
# import modules

# 
# YOUR CODE HERE
# 
```

## Problem 5 (extra)
### (3 points; manually graded)

#### Reload the third dataset from Problem 3

```python
datafile = f'DataFiles/PS1_Prob3_dataset3.txt'
X, y = np.loadtxt(datafile, unpack=True)
m = len(X)
X = X.reshape(m,1); y = y.reshape(m,1)
```

#### Validation curves
Create a validation curve for the polynomial model (for different degrees) where you plot the training score and the validation score as a function of the model complexity. 
- The model complexity is simply the degree of the polynomial.
- Make sure to print (e.g. in the axis label) what "score" that is shown (remember that different measures are being used in the literature).
- For the computation of training and validation scores you should use cross-validation which is more stable than the use of a single, random pair of training and validation sets.

Question to ponder:
- From this curve: Which order polynomial do you think was used when generating the data?

*Hint:* `scikit-learn` has a function `validation_curve` that might be useful. It is instructive to implement the calculation of learning curves yourself, but it is not required for this task.

<!-- grade_id: cell-4c0faa3d3a1db40e, points: 0, grade: false, solution: true -->
```python
# 
# YOUR CODE HERE
# 
```

## Problem 6 (extra)
### (3 points; manually graded)

### Gradient descent methods

#### Generate noisy data with a quadratic feature
This is the same data as in Problem 2.

```python
# Generate noisy data with a quadratic feature
# use the following code:
np.random.seed(42)

# X are picked uniform random [0,2]
X = 2 * np.random.rand(100, 1)
# Linear relation to the predicted value, but with Gaussian noise (mean=0, standard deviation=0.2)
theta_true = [0.25, 1, 0.75]
noise_std = 0.2
y = theta_true[2] * X**2 + theta_true[1] * X + theta_true[0] \
+  np.random.normal(loc=0.0, scale=noise_std, size=(m,1))
```

#### (a) Batch and stochastic gradient descent
Implement both batch and stochastic gradient descent and use these methods to find the best fit parameters of a quadratic model. Make sure that you also save the convergence path, i.e., how the parameters change as a function of iteration number. Concerning batch gradient descent you can use the methods that you implemented in Problem 2 with relevant modifications.
- You might want to tune the learning hyperparameter $\eta$.
- Do 50 epochs for the SGD (each epoch corresponding to using all instances of data once).
- You can start from $\theta=(0,0,0)$.
- Compare with the solution from Problem 2.

<!-- grade_id: bgd, points: 0, grade: false, solution: true -->
```python
# Implement BGD and use it to find the best-fit parameters
#
# At the end, the following array should contain the 
# best-fit parameters: 
# theta_0 (constant term), theta_1 (linear), theta_2 (quadratic)
theta_bgd = np.array([0., 0., 0.]) # Note the order

# 
# YOUR CODE HERE
# 
```

<!-- grade_id: sgd, points: 0, grade: false, solution: true -->
```python
# Implement SGD and use it to find the best-fit parameters
#
# At the end, the following array should contain the 
# best-fit parameters: 
# theta_0 (constant term), theta_1 (linear), theta_2 (quadratic)
theta_sgd = np.array([0., 0., 0.]) # Note the order

# 
# YOUR CODE HERE
# 
```

<!-- grade_id: correct_bgd_sgd, points: 1, grade: true, solution: false -->
```python
# TESTS TO CHECK YOUR CODE
# ---
assert theta_bgd.shape ==(3,)
assert not (theta_bgd==0).any()
assert theta_sgd.shape ==(3,)
assert not (theta_sgd==0).any()

print('Error printouts (if any):')
bgd_rel_residual = \
    (theta_bgd - np.array([0.37086534, 0.70324876, 0.87805028])) / \
     theta_bgd
assert (np.abs(bgd_rel_residual) < 2e-1).all(),\
    print('bgd_rel_residual = \n', bgd_rel_residual)

sgd_rel_residual = \
    (theta_sgd - np.array([0.36208325, 0.70543074, 0.86862094])) / \
     theta_sgd
assert (np.abs(sgd_rel_residual) < 1e-1).all(),\
    print('sgd_rel_residual = \n', sgd_rel_residual)
```

#### (b) Mini-batch gradient descent
Implement mini-batch gradient descent and use this method to find the best fit parameters. Details:
- use a mini-batch size of 20%
- perform 50 epochs 
- hint: the `numpy.random.permutation` function might be useful for creating the mini-batches. 
- Compare with the solutions from BGD and SGD.

<!-- grade_id: mbgd, points: 0, grade: false, solution: true -->
```python
# Implement MBGD and use it to find the best-fit parameters
#
# At the end, the following array should contain the 
# best-fit parameters: 
# theta_0 (constant term), theta_1 (linear), theta_2 (quadratic)
theta_mbgd = np.array([0., 0., 0.]) # Note the order

# 
# YOUR CODE HERE
# 
```

<!-- grade_id: correct_mbgd, points: 1, grade: true, solution: false -->
```python
# TESTS TO CHECK YOUR CODE
# ---
assert theta_mbgd.shape ==(3,)
assert not (theta_mbgd==0).any()

print('Test error printouts (if any):')
mbgd_rel_residual = \
    (theta_mbgd - np.array([0.39386575, 0.7356066 , 0.89338186])) / \
     theta_mbgd
assert (np.abs(mbgd_rel_residual) < 1e-1).all(),\
    print('mbgd_rel_residual = \n', mbgd_rel_residual)
```

#### (c) Convergence
Compare the convergence pattern for the BGD, SGD and MGD, i.e. plot the path towards the optimal set of parameters. 
- You can start from $\theta=(0,0,0)$ with all three GD algorithms.
- Note that the path is in 3D (since there are three parameters in our quadratic model). You should plot the three different 2D projections of the path.
- Indicate also the true optimum (as found by solving the normal equation).

<!-- grade_id: cell-a0ed6a6e81042c13, points: 0, grade: false, solution: true -->
```python
# 
# YOUR CODE HERE
# 
```