## Instructions
- See deadline on the course web page
- This problem set is solved individually. See examination rules on the course web page and the explanation of the examination procedure below.
- The two notebooks for each problem set contain a number of basic and extra problems; you can choose which and how many to work on. The extra problems are usually more challenging.
- Students are allowed to discuss together and help each other when solving the problems. However, every student must understand their submitted solution in the sense that they should be able to explain and discuss them with a peer or with a teacher.
- While discussions with your peers are allowed (and even encouraged), direct plagiarism is not. Every student must reach their own understanding of submitted solutions according to the definition in the previous point.
- The use of coding assistance from code generating artificial intelligence tools is allowed. However, every student must reach their own understanding of submitted solutions (including employed algorithms) according to the definition above.
- Some problems include checkpoints in the form of `assert` statements. These usually check some basic functionality and you should make sure that your code passes these statements without raising an `AssertionError`. 
- Do not use other python modules than the ones included in the `environment.yml` file in the course github repo. 

- **Important:** The grading of problem sets requires **all** of the following actions:
  1. Make sure to always complete **Task 0** in the header part of the notebook and that this part does not raise any `AssertionError`(s).
  1. **Complete** the corresponding questions in Yata for every task that you have completed. This usually involves copying and pasting some code from your solution notebook and passing the code tests. You need to have a green check mark on Yata to get the corresponding points.
  1. **Upload** your solution in the form of your edited version of this Jupyter notebook via the appropriate assignment module in Canvas (separate for basic and extra tasks). It is the code and results in your submitted notebook that is considered to be your hand-in solution.
  1. If selected, be **available for a discussion** of your solution with one of the teachers on the Monday afternoon exercise session directly following the problem set deadline. No extra preparation is needed for these discussions apart from familiarity with your own solution. A list of randomly selected students will be published on the course web page around Monday noon. During the afternoon session that same day, students will be called in the numbered order until the end of the list (or the end of the exercise session). You must inform the responsible teacher as soon as possible following the publication of the student list if you can not be physically present at the exercise session (in which case we will have the discussion on zoom). An oral examination (on all aspects of the course) will be arranged during the exam week for students that do not show up for their discussion slot, or that fail to demonstrate familiarity with their hand-in solutions.

- Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

- Make sure that the **run time is smaller than a few minutes**. If needed you might have to reduce some computational tasks; e.g. by decreasing the number of grid points or sampling steps. Please ask the supervisors if you are uncertain about the run time. 

- Your solutions are usually expected where it says `YOUR CODE HERE` or <font color="red">"PLEASE WRITE YOUR ANSWER HERE"</font>.

### Task 0 
#### (0 points)

By changing the below boolean variable `student_self_assessment` to `True` you attest that:
- All handed in solutions were produced by yourself in the sense that you understand your solutions and should be able to explain and discuss them with a peer or with a teacher.


<!-- grade_id: student-self-assessment, points: 0, grade: false, solution: true -->
```python
student_self_assessment = False

# 
# YOUR CODE HERE
# 
```

<!-- grade_id: cell-795bedd2908899aa, points: 0, grade: true, solution: false -->
```python
assert student_self_assessment == True, 'You must assert the individual solution statements.'
```

# Problem Set 2 (Basic problems)
**Learning from data [TIF285], Chalmers, Fall 2023**

Last revised: 11-Sep-2023 by Christian Forssén [christian.forssen@chalmers.se]

```python
import os
# Data files are stored in
DATA_DIR = "DataFiles/"

if not os.path.exists(DATA_DIR):
    os.makedirs(DATA_DIR)
```

## Problem 1: Coin tossing (2 points)

**Task (see also Yata):** Define a function `bayesian_analysis_coin_flips` (see start code below) that returns the mean, median, and 68%/95% credible intervals of the Bayesian posterior with an input data array of coin flips and using a uniform $[0,1]$ prior for `pH` (probability of heads).

**Task (mainly for testing your code; not on Yata)** Read the data with simulated coin tosses from the file `cointosses.dat`.
Each row corresponds to a single toss: 0=tails; 1=heads

Extract the mean and the 95% credible intervals (Degree-of-belief or DoB intervals) from the first 8 tosses, the first 64 tosses, the first 512 tosses and all 4096 tosses in the data assuming a uniform prior for the probability $p_H$ of obtaining heads in a single toss.

*Hint*: Sample code for computing the DoB interval is available in the demonstration notebook `demo-BayesianBasics.ipynb`.

<!-- grade_id: cell-f27c4016e4570c15, points: 0, grade: false, solution: true -->
```python
# importing modules

%matplotlib inline

import numpy as np
import matplotlib.pyplot as plt
#...

# 
# YOUR CODE HERE
# 
```

```python
# Read data
data = np.loadtxt(f'{DATA_DIR}/PS2_Prob1_cointosses.txt')
```

<!-- grade_id: cell-79bec52ebb8e7d12, points: 0, grade: false, solution: true -->
```python
# Optional. 
# Insert utility code / functions here.

# 
# YOUR CODE HERE
# 
```

<!-- grade_id: cell-538f840bea75b7dd, points: 0, grade: false, solution: true -->
```python
# Define a function that returns the mean, median, and 68%/95% credible intervals 
# of the Bayesian posterior with an input data array of coin flips 
# and using a uniform [0,1] prior for the pH (probability of heads).

def bayesian_analysis_coin_flips(data_coin_tosses):
    """
    Returns various Bayesian analysis results for the given data of coin tosses.
    
    The posterior is p( pH | data, I).
    Assume a uniform p(pH|I) = U[0,1] prior
    
    Args:
        data_coin_tosses: Array of shape (m,) with 'm' independent binary data.
            0 = tails; 1 = heads
            
    Returns:
        (mean, mode, median, dob68, dob95): A tuple with the following elements
            mean: The mean of the posterior distribution (float)
            mode: The mode of the posterior distribution (float)
            median: The median of the posterior distribution (float)
            dob68: A tuple (lo,hi) with the lower and upper limits of the 
                68% degree-of-belief range of the posterior distribution (float,float)
            dob95: A tuple (lo,hi) with the lower and upper limits of the 
                95% degree-of-belief range of the posterior distribution (float,float)
    """
    # 
    # YOUR CODE HERE
    # 
```

<!-- grade_id: cell-a567356cf18d3da5, points: 2, grade: true, solution: false -->
```python
(mean, mode, median, dob68, dob95) = bayesian_analysis_coin_flips(data[:1])
for output in (mean, mode, median, dob68[0], dob95[0]):
    assert output.dtype=='float64', 'Wrong type'
assert len(dob68)==2, 'DoB tuple should be of length 2'
assert len(dob95)==2, 'DoB tuple should be of length 2'
assert np.abs(mean-0.667)<0.001
assert np.abs(mode-1.0)<0.001
assert np.abs(median-0.707)<0.001

# 8 tosses
(mean, mode, median, dob68, dob95) = bayesian_analysis_coin_flips(data[:8])
assert np.abs(mode-0.3750)<0.001
assert np.abs(dob68[0]-0.2469)<0.001
assert np.abs(dob68[1]-0.5538)<0.001
# 1024 tosses
(mean, mode, median, dob68, dob95) = bayesian_analysis_coin_flips(data[:1024])
assert np.abs(median-0.4922)<0.001
# 4096 tosses
(mean, mode, median, dob68, dob95) = bayesian_analysis_coin_flips(data[:4096])
assert np.abs(dob95[0]-0.4727)<0.001
assert np.abs(dob95[1]-0.5033)<0.001
```

Output:
```
TypeError                                 Traceback (most recent call last)
Cell In[6], line 1
----> 1 (mean, mode, median, dob68, dob95) = bayesian_analysis_coin_flips(data[:1])
      2 for output in (mean, mode, median, dob68[0], dob95[0]):
      3     assert output.dtype=='float64', 'Wrong type'

TypeError: cannot unpack non-iterable NoneType object
```

## Problem 2: Bayesian linear regression (3 points)

We will be fitting a linear model (in this case a first order polynomial) to a set data. Our model has two parameters $\vec\theta=[\theta_0,\theta_1]$ (pay attention to the indexing)

$$
y_M(x) = \theta_0 + \theta_1 x 
$$

And our statistical model assumes that errors are independent and identically distributed

$$
y_i = y_M(x_i;\theta) + \varepsilon_i.
$$

Specifically, $\varepsilon_i \sim \mathcal{N}(0, \sigma^2)$ and we assume a fixed standard deviation $\sigma = 40$. 

(Note that the $\varepsilon_i \sim \ldots$ notation means that $\varepsilon_i$ is a draw of a random variable that follows the specified distribution.)

The data is generated from a straight line with intercept = 15. and slope = 1.5 plus random noise.

```python
intercept = 15.
slope = 1.5
theta_true = np.array([intercept, slope])
sigma=40.
```

Read the data from the file `PS2_Prob2_data.txt`.

<!-- grade_id: cell-0609945db7cac2cd, points: 0, grade: false, solution: true -->
```python
# Load the data and plot with fixed error bar 
# Use np.loadtxt() for loading data (the argument 'unpack=True' is useful)
# and plt.errorbar() for plotting data with errorbars

# 
# YOUR CODE HERE
# 
```

### (a) Bayesian linear regression: Implement log prior(s) and likelihood

**Task (see also Yata)** Create functions that return the natural logarithm of the likelihood and two different choices of prior. Since we are seeking a parameter posterior you do not have to use proper normalization of the probability densities. 

Consider two different choices for the prior:
1. A uniform prior:
   $$
   p(\theta_0, \theta_1 | I) \propto \left\{ \begin{array}{ll} 1 & \text{if } -100 \le \theta_0 \le 100 \text{ and } -100 \le \theta_1 \le 100 \\ 0 & \text{else} \end{array}\right.
   $$
2. A uniform prior for the intercept and a symmetry-invariant one for the slope
   $$
   p(\theta_0, \theta_1 | I) \propto \left\{ \begin{array}{ll} \frac{1}{(1+\theta_1^2)^{3/2}} & \text{if } -100 \le \theta_0 \le 100  \\ 0 & \text{else}\end{array}\right.
   $$

<!-- grade_id: cell-85c54c4a189fc811, points: 0, grade: false, solution: true -->
```python
def log_flat_prior(theta):
    '''
    Returns the log uniform prior (-100 <= theta_i <= 100)
    
    Args:
        theta: array of floats with two elements. theta[0]=intercept. theta[1]=slope
        
    Returns:
        logPi: (float) log prior (not necessarily normalized)
    '''
    # 
    # YOUR CODE HERE
    # 
    
def log_symmetric_prior(theta):
    '''
    Returns the log uniform (for theta_0) and symmetric (for theta_1) prior 
    
    (-100 <= theta_0 <= 100; p(theta_1) \propto (1+theta_1^2)^(-3/2))
    
    Args:
        theta: array of floats with two elements. theta[0]=intercept. theta[1]=slope
        
    Returns:
        logPi: (float) log prior (not necessarily normalized)
    '''
    # 
    # YOUR CODE HERE
    # 
```

```python
assert isinstance(log_flat_prior([0.,0.]), (np.floating, float)), 'The output should be a float'
assert isinstance(log_symmetric_prior([0.,0.]), (np.floating, float)), 'The output should be a float'
assert log_flat_prior([0.,0.]) - log_flat_prior([10.,-10.]) == 0., 'The flat prior should be constant in the interval'
assert np.abs(log_symmetric_prior([-10.,1.]) - log_symmetric_prior([20.,2.]) - 1.3744360978112324) <0.001, \
'The log symmetric prior does not evaluate correctly.'
```

<!-- grade_id: cell-958720f73ee427a3, points: 0, grade: false, solution: true -->
```python
def log_likelihood(theta, x, y, dy=sigma):
    '''
    Returns the log likelihood.
    
    Args:
        theta: array of floats with two elements. theta[0]=intercept. theta[1]=slope
        x: data (independent variable). array of floats
        y: data (dependent variable). array of floats
        dy: fixed error (optional; default=sigma defined above), standard deviation of a normal distribution
        
    Returns:
        logL: (float) log likelihood
    '''
    # 
    # YOUR CODE HERE
    # 
```

```python
x_test,y_test = np.loadtxt(f'{DATA_DIR}/PS2_Prob2_data.txt',unpack=True)
assert isinstance(log_likelihood([0,0], x_test, y_test), (np.floating, float)), 'The output should be a float'
assert np.abs(log_likelihood([0.,0.], x_test, y_test) - log_likelihood([10.,1.], x_test, y_test) + 16.588440096874997) <0.001, \
'The log likelihood does not evaluate correctly.'
```

### (b) Plot the posterior for the two different prior choices

**Tasks (see also Yata)**
* Where is the mode of the posterior with these two different priors?
* Plot the joint pdf for the two different prior choices.
* Are the parameters correlated / anti-correlated?

<!-- grade_id: cell-ba8d7bc36ae04469, points: 0, grade: false, solution: true -->
```python
# We'll start by defining a function which takes a two-dimensional grid of likelihoods and 
# returns 1, 2, and 3-sigma contours. This acts by sorting and normalizing the values and then 
# finding the locations of the  0.682 ,  0.952 , and  0.9972  cutoffs:
def contour_levels(grid):
    """Compute 1, 2, 3-sigma contour levels for a gridded 2D posterior"""
    _sorted = np.sort(grid.ravel())[::-1]
    pct = np.cumsum(_sorted) / np.sum(_sorted)
    cutoffs = np.searchsorted(pct, np.array([0.68, 0.95, 0.997]) ** 2)
    return np.sort(_sorted[cutoffs])

# Optional. 
# Insert utility code / functions here.

# 
# YOUR CODE HERE
# 
```

<!-- grade_id: cell-97243b6cddebf189, points: 0, grade: false, solution: true -->
```python
# The dictionary MAP (= maximum a posteriori) should return the mode 
# of the posterior distribution.
# The key is the prior and the value is resulting posterior mode (peak)
# given as theta* = [theta0*, theta1*]
MAP={}
MAP['uniform_prior'] = [0.0, 0.0]
MAP['symmetric_prior'] = [0.0, 0.0]

# 
# YOUR CODE HERE
# 
```

```python
assert np.abs(MAP["uniform_prior"][1] - 1.977) < 0.01
assert np.abs(MAP["symmetric_prior"][0] - 18.725) < 0.01
```

- Plot the joint posterior for the two model parameters for the two different priors.
- Indicate whether the slope and the intercept are correlated or anti-correlated.

<!-- grade_id: cell-80f2e16a2e4b7897, points: 0, grade: false, solution: true -->
```python
# 
# YOUR CODE HERE
# 
```

## Problem 3: MCMC sampling of a Lorentzian pdf using the random walk Metropolis algorithm (2 points)
Note that you must solve this problem if you want to solve (extra) problem 5.

Say that we have some function that tells us the (possibly unnormalized) probability density for a given position in a one-dimensional space. That is, the function is proportional to a univariate pdf. In general we will not be restricted to univariate distributions, but a key feature of the approach that we will implement here is that it can be straightforwardly extended to many dimensions. 

In this task we will assume a known, specific form of this univariate pdf, namely a Lorentzian (Cauchy) distribution, but it might just as well be some very complicated function that can only be evaluated numerically. All that is needed is some function that, for each position in the parameter space, returns a number that is proportial to the probability density.

Let us start by studying the pdf that we will be sampling from using a random walk (that we will set up using the Metropolis algorithm outlined below).

```python
# Modules needed for this exercise
from scipy.stats import norm
from scipy.stats import cauchy
```

```python
# Draw a number of random samples from the standard Cauchy
r = cauchy.rvs(size=1000)
plt.hist(r, density=True, histtype='stepfilled', alpha=0.2, 
         range=(-10,10),bins=21);
```

This histogram corresponds to a finite sample from the pdf of a standard Cauchy (Lorentzian)
$$ 
p(x | \alpha=0, \beta=1) = \frac{1}{\pi(1+x^2)}, 
$$
with its mode and median at $\alpha=0$ and a full-width at half maximum (FWHM) equal to $2\beta = 2$.

Questions to ponder (not part of this problem; answers not requested):
- How does this pdf compare with a standard normal distribution $\mathcal{N}(x;\mu=0,\sigma^2=1)$?
- The Cauchy distribution is often used in statistics as the canonical example of a "pathological" distribution since both its mean value and its variance are undefined. Do you see mathematically why these moments are undefined?

### (a) Construct a Metropolis sampler for univariate distributions

**Task (see also Yata)** First, turn the posterior into a callable function. You should deliberately remove the normalization to make the point that sampling can be made for an unnormalized pdf. Note that we will work directly with the pdf here (not taking the log as in previous examples).

<!-- grade_id: cell-16bc7d26f9be4be2, points: 0, grade: false, solution: true -->
```python
def posterior_function(x, normalized=False):
    '''
    Return the posterior pdf given by a standard Cauchy (Lorentzian).
    
    Args:
        x: position in a one-dimensional space
        normalized: Return a normalized pdf if True (optional, default=False)
    '''
    # 
    # YOUR CODE HERE
    # 

```

**Task (see also Yata)** Now on to constructing the sampler. The code for a MCMC sampler that uses the Metropolis algorithm is enclosed below. However, it misses a few critical ingredients and your task is to add these at the correct places.

1. At first, you have to specify the initial parameter position (that can be randomly chosen), lets fix it to the input argument `start_position`:

```python
current_position = start_position
```

Then, you propose to move (jump) to another position. You can be very dumb or very sophisticated about how you come up with the proposed jump. The Metropolis sampler is rather dumb and just picks a sample from a symmetric proposal distribution (here we again choose a normal distribution) centered around your current position (i.e. `current_position`) with a certain standard deviation (`proposal_width`) that will determine the length scale for proposed jumps 

2. Use `scipy.stats.norm` to create `proposed_position` as a sample from a random variable that is described by a normal distribution with the mean value equal to the `current_position` and a standard deviation that is given by `proposal_width`.

Next, you evaluate whether that new position is a good place to jump to, or not. We quantify this by comparing the probability density at the `proposed_position` with the one for the `current_position`. Usually you would use the logarithms of the probability densities but we omit this here.

3. Compute both `p_current` and `p_proposal`.

Up until now, we essentially have a hill-climbing algorithm that would propose moves in random directions and only accept the step if the `proposed_position` has higher probability density than `current_position`. Eventually we'll get to the mode (at `x = 0` in this case) from where no more moves will be accepted. However, we want to sample a pdf so we'll also have to sometimes accept moves into regions of lower probability. The key trick is to compute the ratio of the two probabilities,

```python
p_accept = p_proposal / p_current
```

and to interpret this ratio as an acceptance probability. Note that the acceptance probability is obtained by dividing the pdf of the proposed parameter setting by the pdf of the current parameter setting. This implies that the probability density does not necessarily need to be normalized, the normalization factor will anyway cancel out. 

You can see that if `p_proposal` is larger, the acceptance probability will be `> 1` and we'll definitely accept the jump. However, if `p_current` is larger, say twice as large, there'll just be a 50% chance of moving to the proposed position. The acceptance of a proposed step will be decided by sampling another random number.

4. Incorporate the acceptance step by comparing `p_accept` to a random number (uniform [0,1]). The `current_position` should be updated if the `accept` variable is `True`. 

Note that the `current_position` is added to our list of parameter samples at the end of the iteration, regardless of it being a new position or not.

This simple procedure gives us samples from the pdf.

The code below also prints detailed information if the optional keyword argument `verbose=True`.

<!-- grade_id: cell-30d4e414d6f57b40, points: 0, grade: false, solution: true -->
```python
def sampler(posterior_func, no_of_samples=4, start_position=.5, 
            proposal_width=1., verbose=False):
    '''
    Simple (but incomplete) Metropolis sampler function.
    
    Args:
        posterior_func: Function that takes a single positional argument and returns 
            the (possibly unnormalized) pdf value.
        no_of_samples: (integer) Number of samples that will be returned (excluding the start position). 
            (default=4)
        start_position: (float) Start position. (default=0.5)
        proposal_width: (float) Width of Gaussian proposal distribution. (default=1.)
        verbose: (Boolean) Verbose output (default=False)
        
    Returns:
        samples: Array of floats. Length = no_of_samples+1
    '''
    # Verbose printing
    if verbose:
        assert no_of_samples < 11, "Too many samples for printing"
        print(f'Step    Trace  Proposed  Met.ratio  Accept')
        
    # 
    # YOUR CODE HERE
    # 
    
    samples = [current_position]
    for i in range(no_of_samples-1):
        # 
        # YOUR CODE HERE
        # 

        # Compute posteriors of current and proposed position       
        p_current = posterior_func(current_position)
        p_proposal = posterior_func(proposed_position) 
        
        # 
        # YOUR CODE HERE
        # 
        
        # Verbose
        if verbose:
            assert no_of_samples < 11, "Too many samples for visualization"
            print(f'{i:>4}  {current_position:7.3f}   {proposed_position:7.3f}    {min(1.,p_accept):5.3f}     {accept}')
        
        # Possibly update position
        # 
        # YOUR CODE HERE
        # 
        
        samples.append(current_position)
        
    return np.array(samples)
```

The following code cell performs some first tests of the functionality of your sampler. It also makes a call to the sampler with the `verbose` option turned on. 
- How many proposed steps are accepted?
- What can we say about the acceptance of proposed steps that would take us closer to the mode of the distribution?
- And how about those that are proposed away from the mode?
You can re-evaluate this cell multiple times to observe several chains.

<!-- grade_id: cell-2ae464bbbc58c95f, points: 2, grade: true, solution: false -->
```python
test_samples = sampler(posterior_function, no_of_samples=100)
assert test_samples[0]==0.5, 'The first position sample should be the default start position = 0.5'
assert len(test_samples)==100, 'The sampler chain should contain the specified number of samples'

sampler(posterior_function, start_position=0., no_of_samples=10, verbose=True)
```

### (b) Sampling the posterior
**Task (see also Yata)** Draw 100,000 samples from your sampler and plot:
* The trace (i.e. the sequence of draws of your single parameter x)
* A normalized histogram of the samples compared to the true posterior pdf (normalized).

```python
try:
    samples = sampler(posterior_function, no_of_samples=100000, start_position=1.)
except:
    samples=None
    print('The method "sampler" must be defined and working as expected.')
```

<!-- grade_id: cell-17e4557f37667278, points: 0, grade: false, solution: true -->
```python
# Plotting commands here
#
# 
# YOUR CODE HERE
# 
```

<!-- grade_id: cell-02cd9d30ab8b4c8c, points: 1, grade: true, solution: false -->
```python
# If the code works as expected your trace of samples should have a median value close to 0.0
assert np.abs(np.median(samples))<0.2, f'The median of the samples is {np.median(samples):.1}. It should be close to the median of the posterior pdf (=0.0)'
# Furthermore, the Cauchy distribution has heavy tails and your chain will surely have ventured out to rather large distances 
# (providing samples from the tails)
assert np.abs(samples).max()>20, f'The furthest sample is {np.abs(samples).max():.1f}. Too small variations'
```

## Problem 4: Signal and background (3 points)
This problem is more demanding than the other basic problems. Note that you must solve this problem if you want to solve (extra) problem 5.

The goal of this problem is to estimate the amplitude of a signal when there is a background.  We'll take a limiting case where the background is flat, so it is completely specified by its magnitude $B > 0$, and the signal is known to be a Gaussian with unknown amplitude $A$ but known position (mean) and width (standard deviation). 

The measurements will be integer numbers of counts $\{N_k\}$ in well-defined (equally spaced) bins $\{x_k\}$. The index $k$ runs over integers labeling the bins.

We can imagine three different goals of the data analysis:
- Finding $A$ and $B$ given $\{N_k\}$.
- Finding $A$ (we do not care about $B$).
- Finding $B$ (we do not care about $A$).

In all cases we consider the bin sizes and the signal shape (including its mean position and width) as known information.

Our statistical model includes the true signal plus a constant background. The signal and the background magnitudes are the unknown parameters while the other parameters dictating the signal (width $w$ and mean $x_0$ of the Gaussian) are known and fixed:

$$
   D_k = n_0 \left[ A e^{-(x_k-x_0)^2/2 w^2} + B \right]
$$

Here $n_0$ is a constant that scales with measurement time.  Note that $D_k$ is not an integer in general, unlike $N_k$.

<!-- grade_id: cell-ade2557e6088a291, points: 0, grade: false, solution: true -->
```python
# import statements.
# We use pickle to save and load a python dictionary
import pickle
# factorial from the math module is useful. You might consider other modules as well.
from math import factorial

# import additional modules as needed
# 
# YOUR CODE HERE
# 
```

```python
# This function generates data according to the statistical model
A_true = 1.
B_true = 2.

def exact_data(A, B, n_0, x_k, x_0=0., width=np.sqrt(5.)):
    """
    Return the exact signal plus background.  The overall scale is n_0,
    which is determined by how long counts are collected. 
    The default signal position and width are 0.0 and sqrt(5), respectively (in some  irrelevant units).
    """
    return n_0 * (A * np.exp(-(x_k - x_0)**2/(2.*width**2)) + B)
```

#### Poisson distribution
We are imagining a counting experiment, so the statistics of the counts we record will follow a Poisson distribution. It might be an interesting exercise to derive why this distribution is expected for a counting experiment. 

The Poisson discrete random variable from scipy.stats is defined by (see [documentation](https://docs.scipy.org/doc/scipy/reference/generated/scipy.stats.poisson.html))

$$
p(k \mid \mu) = \frac{\mu^k e^{-\mu}}{k!} \quad \mbox{for }k\geq 0 \;.
$$

where $k$ is an integer and $\mu$ is called the shape parameter. The mean and variance of this distribution are both equal to $\mu$. Sivia and Gregory each use a different notation for for this distribution, which means you need to be flexible. 

For convenience, we'll define our own version in this notebook:

$$
p(N \mid D) = \frac{D^N e^{-D}}{N!} \quad \mbox{for }N\geq 0 \;.
$$

where $N$ is an integer.

```python
# make a dataset for exploring
def make_dataset(A_true, B_true, width, x_0, databins=15, delta_x=1, D_max=100):
    """
    Create a data set based on the number of bins (databins), the spacing
    of bins (delta_x), and the maximum we want the exact result to have
    (D_max, this fixes the n_0 parameter).
    
    Return arrays for the x points (xk_pts), the corresponding values of the
    exact signal plus background in those bins (Dk_pts), the measured values
    in those bins (Nk_pts, integers drawn from a Poisson distribution), the 
    maximum extent of the bins (x_max) and n_0.
    """
    # set up evenly spaced bins, centered on x_0
    x_max = x_0 + delta_x * (databins-1)/2
    xk_pts = np.arange(-x_max, x_max + delta_x, delta_x, dtype=int)
    
    # scale n_0 so maximum of the "true" signal plus background is D_max
    n_0 = D_max / (A_true + B_true)  
    Dk_pts = exact_data(A_true, B_true, n_0, xk_pts, width=width)
    
    # sample for each k to determine the measured N_k
    Nk_pts = [stats.poisson.rvs(mu=Dk) for Dk in Dk_pts]
    
    return xk_pts, Dk_pts, Nk_pts, x_max, n_0
```

#### Plot the signal and the data (these tasks are not graded but will help you to understand the problem)
* Make a plot of the true signal plus background we are trying to deduce. Use $A_\mathrm{true}=1$ and $B_\mathrm{true}=2$ and the signal position (mean) $x_0=0$ and width (standard deviation)  $w=\sqrt{5}$.

We consider what happens for fixed signal and background but changing the experimental conditions specified by `D_max` and `databins` (we'll keep `delta_x` fixed to 1). In all cases the bins are symmetric around $x=0$.

The pickle file that is loaded in the cell below contains data from four differently designed counting experiments.:
1. Baseline case: 15 bins and maximum expection of 100 counts per bin.
1. Low statistics case: 15 bins and maximum expection of only 10 counts per bin.
1. Greater range case: 31 bins (with fixed bin width) and maximum expection of 50 counts per bin to give approximately the same total number of counts as in baseline case.
1. Smaller range case: 7 bins (with fixed bin width) and maximum expection of 200 counts per bin to give approximately the same total number of counts as in baseline case.
 
* Make four subplots that correspond to the data from the different experiments.

<!-- grade_id: cell-b300bc6fd3ee1f07, points: 0, grade: false, solution: true -->
```python
# Plotting the signal and background
#
# 
# YOUR CODE HERE
# 
import pickle

# The data has been generated already and will be loaded from a pickle file.
# It is a dictionary with four keys corresponding to the four cases, and each value
# corresponding to xk_pts, Dk_pts, Nk_pts, x_max, n_0
with open(f'{DATA_DIR}/PS2_Prob4_data.pickle','rb') as f:
    data = pickle.load(f)
    print('Loaded "data" dictionary from file.')
    print('Extract data with:')
    print('xk_pts, Dk_pts, Nk_pts, x_max, n_0 = data[case]')
    print('where the key "case" is one of:')
    cases = data.keys()
    print(cases)

# Plotting the data for the four cases
#
# 
# YOUR CODE HERE
# 
```

Output:
```
Loaded "data" dictionary from file.
Extract data with:
xk_pts, Dk_pts, Nk_pts, x_max, n_0 = data[case]
where the key "case" is one of:
dict_keys(['Baseline', 'Low Statistics', 'Greater Range', 'Smaller Range'])
```

### (a) Compute and plot the joint posterior on a grid
**Tasks (see also Yata)**
* Implement functions for the (log) likelihood and for a uniform (log) prior. Let's use a uniform prior for $0 \le A \le 5$ and $0 \le B \le 5$.
* Evaluate the log-posterior on a grid and then: 
  - Plot the joint posterior pdf for $A$ and $B$ for the four cases.
  - Plot the marginalized posterior pdf for the signal amplitude $A$ for the four cases.
  
  Use the same axis scales for all four cases such that the precision of the inference can be compared.

<!-- grade_id: cell-a3772dfb9786e84d, points: 0, grade: false, solution: true -->
```python
# Define the pdfs and combine with Bayes' theorem.

def log_prior(A, B):
    """
    Log prior .
    
    We take a uniform (flat) prior with large enough
    maximums but, more importantly, require positive values for A and B.
    """
    A_max = 5.
    B_max = 5.
    # flat prior 
    if np.logical_and(A <= A_max, B <= B_max).all(): 
        return np.log(1/(A_max * B_max))
    else:
        return -np.inf


def log_likelihood(A, B, xk_pts, Nk_pts, n_0):
    """Log likelihood for data Nk_pts given A and B"""
    # 
    # YOUR CODE HERE
    # 
    
def log_posterior(A, B, xk_pts, Nk_pts, n_0):
    """Log posterior for data Nk_pts given A and B"""
    return log_prior(A, B) + log_likelihood(A, B, xk_pts, Nk_pts, n_0)

# Other utility code can be put here (if needed)
#
# 
# YOUR CODE HERE
# 
```

<!-- grade_id: cell-cbb85be603215ec4, points: 0, grade: false, solution: true -->
```python
# Optional.
# Code to find contour levels of gridded 2D posterior.

def find_contour_levels(grid):
    """Compute 1, 2, 3-sigma contour levels for a gridded 2D posterior
       Note: taken from BayesianAstronomy but may not work here.
    """
    sorted_ = np.sort(grid.ravel())[::-1]
    pct = np.cumsum(sorted_) / np.sum(sorted_)
    cutoffs = np.searchsorted(pct, np.array([0.68, 0.95, 0.997]) ** 2)
    return np.sort(sorted_[cutoffs])

# 
# YOUR CODE HERE
# 
```

### (b) Design of experiment
**Tasks (see also Yata)** 
Answer the following questions. Some of them are repeated on Yata.
  1. Can you understand why the signal and background amplitudes are anticorrelated? And why the (anti)correlation seems to be stronger in one of the cases? 
  1. What are your conclusions for how to design the experiment given limited resources? 
    - In particular, given that you wanted to be able to distinguish between signal amplitude and background, would it then be better to have many counts in few bins, or the same total amount of counts spread over a wider interval? 

<!-- grade_id: cell-620bd3518d4f7363, points: 1, grade: true, solution: true -->
```python
# 
# YOUR CODE HERE
# 
```

```python

```